---
title: 'Chapter Title Here'
description: 'Chapter description goes here.'
free_preview: true
---

## Example coding exercise

```yaml
type: NormalExercise
key: e8c1edbe67
lang: python
xp: 100
skills: 2
```

This is an example exercise. This content is meant to test the way DataCamp courses work.

`@instructions`
This is the instructions which should be at least 30 characters in length.

Objective:
Print Hello world.

`@hint`
It's not hi mom. Please be sure you read the instructions.

`@pre_exercise_code`
```{python}

```

`@sample_code`
```{python}
# complete the code
print("your text here")
```

`@solution`
```{python}
Hello world.
```

`@sct`
```{python}
Ex().has_output('Hello')
```

---

## Insert exercise title here

```yaml
type: TabExercise
key: 21ad3779c2
xp: 100
```



`@pre_exercise_code`
```{python}

```
